Sex–specific Single Transcript Level Atlas of Vasopressin and its Receptor (AVPR1a) in the Mouse Brain

Author response:

The following is the authors’ response to the previous reviews

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The authors have adequately addressed all of my concerns. I have no further questions or concerns.

We thank the Reviewer #1.

Reviewer #2 (Recommendations for the authors):

We thank the Reviewer #2 for thoughtful recommendations.

(1) Figure 1A, 1B, 2B, 2C, etc.: The Y-axis label is confusing. I assume the intention was to make big numbers small by dividing by 1000. The comma makes the label confusing. Perhaps, make the label more "mathematical" as in "Avp density ((transcript/µm2) * 10-3)" or rearrange the math to be clearer as in "Avp density (transcript/1000 per µm2)".

Great suggestion and done exactly as suggested in Figures 1, 2 and 4.

(2) Figure 1B and 1C: The figure and legend do not match up. Either switch the figures or the legends. Currently, legend 1B describes image 1C.

Agreed and done as suggested.

(3) Figure 2A is broken up into separate pages/panels. It could be integrated better or separated to make A and B, then shift B and C to C and D.

Great suggestion and we have done exactly as suggested.

(4) Figure 2 legend: I recommend putting the scale bar info with (A) rather than at the end. The stars used in the figure are not explained in the legend.

Good points. We have made all necessary changes as suggested.

(5) Supplementary Figure 1B: The legend states that the data are the number of transcript-containing cells, but the figure states transcript number.

We thank the Reviewer for pointing out this typo. We corrected all graph legends in the Supplementary Figure 1.

Changes from the earlier version to the revised version

Abstract

Vasopressin (AVP), a nonapeptide synthesized predominantly by magnocellular hypothalamic neurons, is conveyed to the posterior pituitary via the pituitary stalk, where AVP is secreted into the circulation. Known to regulate blood pressure and water homeostasis, it also modulates diverse social behaviors, such as pair–bonding, social recognition and cognition in mammals including humans. Importantly, AVP modulates social behaviors in a sex–specific manner, perhaps, due to sex differences in the distribution in the brain of AVP and its main receptor AVPR1a. There is a corpus of integrative studies for the expression of AVP and AVPR1a in various brain regions, and their functions in modulating central and peripheral actions. In order to purposefully address sexually dimorphic and novel roles of AVP on central and peripheral functions through its AVPR1a, we utilized RNAscope to map Avp and Avpr1a single transcript expression in the mouse brain. As the most comprehensive atlas of AVP and AVPR1a in the mouse brain, this compendium highlights the importance of newly identified AVP/AVPR1a neuronal nodes that may stimulate further functional studies.

Introduction

Vasopressin (AVP), a nonapeptide synthesized primarily by magnocellular neurons of the paraventricular nucleus (PVH) and supraoptic nucleus (SO) of the hypothalamus, is conveyed along axons of magnocellular neurons via the pituitary stalk to the posterior pituitary, where AVP is released into circulation to exert hormonal functions (1, 2). Involved in the regulation of blood pressure and water balance (3, 4), AVP also regulates diverse social behaviors, such as pair–bonding, social recognition and cognition in all mammals, including humans (5–9). It should be noted that similar to OXT, AVP is evolutionarily conserved across invertebrates and vertebrate taxa, differing from OXT by just only two amino acids. Importantly, AVP modulates social behaviors in sex–specific manner, perhaps due to sex differences in AVP and its receptor expression in the brain (10–15).

AVP receptors are G protein–coupled receptors consisting of two major subtypes, the V1 receptor (AVPR1) and V2 receptor (AVPR2) (16). AVPR1 has two subtypes, namely AVPR1a and AVPR1b that mediate the effects of AVP on social behaviors. Avpr1a will be the overarching focus of this study given its abundance and ubiquity in brain regions (9). In fact, differences in AVPR1a genetic variability and expression patterns determine specific social phenotypes (17–21). There is evidence that AVPR1a is involved in maternal care, pair–bonding, behavioral aggression, anxiety, social recognition and social play (6, 9, 19, 22–25). The sex–dependent distribution of Avpr1a across the brain in different species provides a proxy for the distribution of AVP binding, and therefore, provides further evidence for central AVP neural nodes of physiologic relevance.

Blocking AVPR1a inhibits social recognition in the rat, while AVPR1a knockout mice fail to display social recognition (15, 22, 26). These effects of AVP in enhancing social recognition is mediated via AVPR1a in the lateral septum (15, 26). Sex differences in Avpr1a binding densities have been described in several brain sites of Wistar rats. Namely, males display higher AVPR1a binding densities in the following forebrain areas: somatosensory and piriform cortex, medial posterior bed nucleus of the stria terminalis (BNST), nucleus of the lateral olfactory tract (LOT), anteroventral thalamic nucleus (VA), tuberal lateral hypothalamus (LH), stigmoid hypothalamus (Stg), and dentate gyrus (DG) (16). In male prairie voles, central AVP infusion facilitates selective aggression associated with pair bond formation and partner preferences, and the Avpr1a antagonist 1-(β-mercapto-β, β-cyclopentamethylene propionic acid) does not seem to inhibit aggression (10). In contrast, central AVP infusion induces a partner preference in female prairie voles (11). Comparative analysis of AVPR1a distribution has revealed higher densities of AVPR1a binding in the ventral pallidum, amygdala, and thalamus of prairie voles than that of meadow or montane voles (27, 28). It has also been reported that AVPR1a antagonism specifically in the ventral pallidum prevents mating–induced partner preferences in male prairie voles (29).

Although osmotically stimulated AVP release with short latency and duration from hypothalamic PVH and SO magnocellular axon terminals occurs in the pituitary (30, 31), it is also released locally from somata and dendrites in the SO with a longer delay responding to osmotic challenge (32, 33). Such local release is likely to facilitate autocrine and/or paracrine regulation of SO–magnocellular neuronal activity and inhibit further systemic AVP output (33, 34). Of note, somatodendritic AVP release in response to direct hypertonic stimulation is attenuated by V1/V2 receptor antagonism, implying that AVP may facilitate its own release by acting on autoreceptors within magnocellular neurons of the SO (35). The importance of autofacilitation to local AVP release may lay in fine–tuned regulation of AVP actions towards physiologic demands. Alternatively, AVP could putatively diffuse over longer distances to bind to adjacent receptors. The relative contribution of AVP autoreceptor subtypes, including AVPR1a, to this phenomenon awaits further clarification.

Current studies implicate posterior pituitary hormones, traditionally thought of as master regulators of a single physiological target, in the control of multiple bodily systems, either directly or via their receptors in the brain (23, 36–39). Non–traditional actions of AVP include its ability to affect skeletal homeostasis, wherein it negatively regulates osteoblasts and stimulates osteoclasts. This explains the bone loss that accompanies low blood sodium levels in patients with high AVP levels (40). Furthermore, we have shown that AVP (via AVPR1a) and oxytocin (via OXTR) have opposing skeletal actions—effects that may relate to the pathogenesis of bone loss in chronic hyponatremia, and pregnancy and lactation, respectively (40–43).

Detecting specific AVPR1a and AVPR1b in the brain has had limitations for a long time due to the availability of only nonselective radioligands, such as tritium labeled (3H) AVP ligands, which bind to both receptors (44). Although there is a large body of integrative studies for the expression of AVP and AVPR1a in various brain regions, and their functions in regulating central and peripheral actions, there remains the need for detailed, sex–specific mapping of the AVP/AVPR1a neuronal nodes in the brain. We utilized RNAscope—a technology that detects single RNA transcripts—to create a comprehensive atlas of AVP and AVPR1a in the mouse brain. It may seem somewhat remarkable that newly discovered brain areas for receptors of such evolutionarily conserved and well–studied hormones AVP and OXT are currently emerging, with inferences to novel functions. We believe that this atlas of AVP and its AVPR1a in concrete brain sites at a single transcript level should provide a resource to neuroscientists to deepen our understanding of classical and novel central and peripheral functions of AVP by interrogating AVPR1a site–specifically.

Results

AVP receptors are G protein-coupled receptors consisting of two major subtypes, the AVPR1 and AVPR2 (16). In turn, AVPR1 is divided into AVPR1a and AVPR1b receptor subtypes that mediate the effects of AVP in the brain on social behaviors. In this study, we have provided not only a distribution mapping of AVP and AVPR1a in the brain, but also assessed sex differences in their expression by RNAscope. Allowing the detection of single transcripts, RNAscope uses ∼20 pairs of transcript–specific double Z–probes to hybridize 10–µm–thick whole brain sections. Preamplifiers first hybridize to the ∼28–bp binding site formed by each double Z–probe; amplifiers then bind to the multiple binding sites on each preamplifier; and finally, labeled probes containing a chromogenic enzyme bind to multiple sites of each amplifier.

RNAscope data were quantified on every tenth section of the whole brains from coded 3 female and 3 male mice. For simplicity and clarity in the graphs, a scatter plot has been shown for 3 nuclei, sub–nuclei and regions with the highest AVP and AVPR1a transcript densities as well as their absolute transcript counts. Each section was viewed and analyzed using CaseViewer 2.4 (3DHISTECH, Budapest, Hungary) and QuPath v.0.2.3 (University of Edinburgh, UK). The Atlas for the Mouse Brain in Stereotaxic Coordinates (45) was used to identify every nucleus, sub–nucleus or region, which was followed by manual counting of Avp and Avpr1a transcripts by two independent observers (V.R. and A.G.) in every tenth section using a tag feature. Receptor density was calculated by dividing the absolute receptor number by the total area (µm2, ImageJ) of every nucleus, sub–nucleus or region. The highest Avp and Avpr1a values in the brain regions are presented as means ± SE and compared with those of the opposite sex. Photomicrographs were prepared using Photoshop CS5 (Adobe Systems) only to adjust brightness, contrast and sharpness, to remove artifacts (i.e., obscuring bubbles), and to make composite plates.

RNAscope revealed Avp expression in the hypothalamus, forebrain, hippocampus and cortex of both female and male mice (Fig. 1A), however, Avp expression in the 3rd ventricular region and thalamus was found only in the female mouse (Fig. 1B). Whereas the numbers of Avp–expressing cells were greater in females compared with males in the hypothalamus (607 vs. 471), 3rd ventricular region (7 vs. 0) and thalamus (2 vs. 0), those numbers were lower in the hippocampus (58 vs. 118), forebrain (50 vs. 77) and cortex (5 vs. 6) [see Appendix 1 for Glossary and Supplementary Fig. 1 for raw count graphs].

Sex–specific Avp transcript density in the brain.

Earlier version: (A) Avp transcript density in main brain regions detected by RNAscope. Female 3V region and hypothalamus and male hypothalamus had the highest Avp transcript density. (B) Representative micrographs of some of the hypothalamic and forebrain regions with highest Avp expression. Scale bar: 20 µm. (C) Avp transcript density in nuclei, sub– nuclei and regions of the hypothalamus and forebrain. (D) Novel Avp transcripts found in nuclei, sub–nuclei and regions of the hippocampus, thalamus and cerebral cortex. Scale bar: 10 µm. Note, that Avp transcripts were found only in the female 3V ependymal layer and thalamus. N=3 mice per sex.

Revised version: (A) Avp transcript density in main brain regions detected by RNAscope. Female 3V region and hypothalamus and male hypothalamus had the highest Avp transcript density. (B) Avp transcript density in nuclei, sub–nuclei and regions of the hypothalamus and forebrain (C) Representative micrographs of some of the hypothalamic and forebrain regions with highest Avp expression. Scale bar: 20 µm. (D) Novel Avp transcripts found in nuclei, sub–nuclei and regions of the hippocampus, thalamus and cerebral cortex. Scale bar: 10 µm. Note, that Avp transcripts were found only in the female 3V ependymal layer and thalamus. N=3 mice per sex. ***P<0.001 and **P<0.01.

The highest Avp transcript densities were detected in following brain nuclei, sub–nuclei and regions of female and male mice, respectively: ventricular region—3V only for females, hypothalamus—SChVL and PaDC, forebrain—MPOM and VLPO, hippocampus—Py and GrDG, thalamus—ic only for females and cerebral cortex—Pir for both (Fig. 1B). Representative micrographs of some of the hypothalamic and forebrain regions with highest Avp expression are shown in Fig. 1C.

We found that Avpr1a transcript expression in several brain nuclei, for example, the MS of the forebrain, medullary IOBe, IRt, LRt (Supplementary Fig. 2B), displayed individual patterns of expression vs. a more ubiquitous and even expression noted in most of the other brain areas, suggesting brain–site–specific functional diversity and context–selective regulation of Avpr1a–mediated signaling. We report the expression of the Avpr1a in 398 and 375 brain nuclei, sub–nuclei and regions of the female and male mice, respectively. Avpr1a transcripts were detected bilaterally, with no apparent ipsilateral domination. Probe specificity was established by a positive signal in renal tubules of the kidney (positive control) with an absent signal in the FrA of the frontal cortex (negative control) (Fig. 2A). Representative micrographs of sex–specific medullary CC and hypothalamic Arc with highest Avpr1a expression also are shown in Fig. 2A.

Sex–specific Avpr1a transcript density in the brain.

Earlier version: (A) Probe specificity was established by a positive signal in renal tubules of the kidney (positive control) with an absent signal in the frontal association cortex (FrA) (negative control). Also shown are representative micrographs of medullary central canal (CC) and hypothalamic arcuate nucleus (Arc). (B) Avpr1a transcript density in main brain regions detected by RNAscope. (C) Avpr1a transcript density in nuclei, sub–nuclei and regions of the ventricular regions, hypothalamus, medulla, midbrain and pons, forebrain, olfactory bulb, hippocampus, thalamus, cortex and cerebellum. N=3 mice per sex. Scale bar: 25 µm (controls) and 20 µm in (A).

Revised version: (A) Probe specificity was established by a positive signal in renal tubules of the kidney (positive control) with an absent signal in the frontal association cortex (FrA) (negative control). Sale bar: 20 µm. (B) Representative micrographs of medullary central canal (CC) and hypothalamic arcuate nucleus (Arc). (C) Avpr1a transcript density in main brain regions detected by RNAscope. (D) Avpr1a transcript density in nuclei, sub–nuclei and regions of the ventricular regions, hypothalamus, medulla, midbrain and pons, forebrain, olfactory bulb, hippocampus, thalamus, cortex and cerebellum. N=3 mice per sex. Scale bar: 25 µm (controls). ***P<0.001, **P<0.01 and *P<0.05.

In the female, total Avpr1a transcript numbers were the highest in the medulla followed, in descending order, by the hypothalamus, cortex, midbrain and pons, forebrain, thalamus, cerebellum, hippocampus, olfactory bulb and ventricular regions (Supplementary Fig. 2A). In the male, Avpr1a transcripts were the highest in the forebrain, followed by the medulla, midbrain and pons, hypothalamus, cortex, hippocampus, thalamus, olfactory bulb, cerebellum and ventricular regions (Supplementary Fig. 2).

Using the RNAscope dataset we further calculated Avpr1a density in all brain divisions (Fig. 2B), nuclei, sub–nuclei and regions (Fig. 2C). Highest Avpr1a densities in the female and male mice, respectively, were noted in nuclei, sub–nuclei and regions as follows: ventricular regions—CC ependymal region for both with 2.92–fold greater in the male; hypothalamus—Arc for both; medulla—IOBe and InM; midbrain and pons—vtgx and DRI; forebrain—AC and LSI; olfactory bulb—Mi and Tu; hippocampus—TS and df; thalamus—PV and Xi; cortex—MPtA and AID and cerebellum—1Cb for both with 2.58-fold greater in the female (Fig. 2C).

In addition, RNAscope analysis revealed that various brain nuclei, sub–nuclei and regions in both female and male mice co–localized Avp and Avpr1a transcripts. Avpr1a to Avp ratios within the same brain nucleus, sub–nucleus and region in both sexes are demonstrated in Figs. 3A, B. Finally, RNAscope showed Avp and Avpr1a expression in the posterior pituitary lobe with Avp and Avpr1a transcript densities that were higher in male compared with female mice, however, without statistical significance (Fig. 4).

Avp and Avpr1a co-localization in the brain.

We found that various nuclei and sub–nuclei exhibited Avp and Avpr1a co-localization in the brain of both sexes. (A) Female and (B) male Avpr1a to Avp ratios in the cortex, forebrain, hippocampus, 3rd ventricular region, thalamus and hypothalamus.

Avp and Avpr1a transcript densities in the pituitary gland.

RNAscope revealed Avp and Avpr1a expression in the posterior pituitary lobe with Avp and Avpr1a transcript densities that were higher in male compared with female mice. N=3 mice per sex.

Discussion

Here, we attempted to integrate previous information on sex–specific AVP and its AVPR1a expression in the murine brain. AVPR1a is the most abundant and wide spread receptor in the brain (9) that plays a dominant role in regulating behavior. In addition, we focused towards paradigm–shifting non–traditional roles of central AVP signaling in light of newly discovered AVPR1a. We report AVP expression in 41 female and 13 male brain nuclei, sub–nuclei and regions. Moreover, we identified abundant AVPR1a expression in 398 female and 375 male brain nuclei, sub–nuclei and regions. Therefore, this report is the most exhaustive atlas of brain Avp and Avpr1a expression at the single transcript level.

It has been reported that AVP synthesis and AVP fiber projections are sexually dimorphic in specific brain sites [for review, see: (16)]. To our knowledge, the first discovery of the sexually dimorphic nature of AVP in the rat brain was made by De Vries et al. (46). That is, males displayed more AVP–immunoreactive fibers in the lateral septum and lateral habenular nucleus over females (46). Surprisingly, sex differences in AVP fiber density in the LS and medial amygdala (MeA) originate from the BNST, given only lesions to the BNST, but not the PVH, result in decreased AVP fiber density in the LS (47, 48). In adult rats, AVP fiber density from the BNST and MeA are dependent on circulating gonadal hormones, as gonadectomy eliminates AVP expression and hormone replacement restores AVP fiber network (49, 50). Nonetheless, gonadal hormones appear to only partially explain sex difference in AVP expression because both females and males, exposed to a similar gonadal steroid hormone regime, still differ sexually (51, 52).

Magnocellular neurons of the PVN, SO and SCh of the hypothalamus are the predominant source of AVP synthesis. Hypothalamic AVP synthesis of most rodent species is similar between males and females in the PVH and SO of mice (53, 54); PVH, SO and SCh of voles (55, 56); PVH, MPOA, LH and AHA of Mongolian gerbils and Chinese striped hamsters (57) [for review, see: (16)]. In concordance with these reports, we also find no sex difference in Avp synthesis, as evidenced by similar numbers of Avp–expressing neurons, in the PVH and SO of mice; however, we did note sex differences in Avp expression density in specific hypothalamic nuclei. Notably, Avp expression density was higher in several PVH (PaLM and PaMM) and suprachiasmatic (SChVL and RCh) subnuclei of female compared to male mice. No sex differences in SO–Avp expression density were found.

Furthermore, Avpr1a transcript density was highest in the arcuate nucleus (Arc) and retrochiasmatic sub–nucleus (RCh) in female compared with Arc and suprachiasmatic nucleus (SCh) of male mice. Avpr1a expression in the Arc has previously been reported (58), however, its role was unclear until a recent report demonstrating a critical involvement of Arc–NPY in the regulation of fluid homeostasis and the induction of salt water–induced hypertension through AVP modulation in the SO (59); the latter receives direct projections from the Arc (60). In contrast, it is plausible, but, by no means proven, that PVH– and/or SO–AVP may modulate Arc anorexigenic neurons to inhibit ingestive behavior, which has previously been shown with PVH oxytocinergic neurons (61). Indeed, increasing evidence suggests that AVP reduces feeding in mammals (62, 63).

Earlier version: It has been reported that in the rat, AVP is an important output of the SCh targeting AVP cells in other hypothalamic areas—its release into the CSF peaks in the early morning and declines later in a day (64). Specifically, SCh–AVP secreted during late sleep activates osmosensory afferents to AVP neurons in the SO and organum vasculosum of the lamina terminalis (65, 66). Similar to rodents, studies in humans also determined that the main AVP projections from the SCh target the anteroventral hypothalamic area, sub–PVH as well as ventral parts of the PVH and DMH––a remarkable evolutionary conservation of SCh innervation from rodent to human (67, 68).

Revised version: It has been reported that in the rat, AVP is an important output of the SCh targeting AVP cells in other hypothalamic areas—its release into the CSF peaks in the early morning and declines later in a day (64). Specifically, SCh–AVP secreted during late sleep activates osmosensory afferents to AVP neurons in the SO and organum vasculosum of the lamina terminalis (65, 66). Similar to rodents, studies in humans also determined that the main AVP projections from the SCh target the anteroventral hypothalamic area, sub–PVH as well as ventral parts of the PVH and DMH––a remarkable evolutionary conservation of SCh innervation from rodent to human (67, 68). The fact that the SCh is another brain nucleus with high AVP and AVPR1a expression density (greater in males vs. females), accentuates an important role of SCh–AVP in circadian rhythmicity, notably impacting neuroendocrine day/night rhythms, feeding timing, period, precision, and synchronization of SCh neurons (64, 69, 70).

Removed in the revised version: The fact that the SCh is another brain nucleus with high AVP and AVPR1a expression density (greater in males vs. females), accentuates an important role of SCh–AVP in circadian rhythmicity, notably impacting neuroendocrine day/night rhythms, feeding timing, period, precision, and synchronization of SCh neurons (64, 69, 70).

In the hindbrain, the highest Avpr1a transcript density was noted in the inferior olive, beta subnucleus (IOBe) of female and intermedius nucleus (InM) of male mice. It has been reported that AVP fibers are apparent in the hindbrain, such as the parabrachial nucleus, locus coeruleus, and near inferior olive nuclei (71). In this regard, Avpr1a mRNA expression has been noted in the inferior olive (58). Given this nucleus has been implicated in various functions including learning and timing of movements, it is possible that AVPR1a in the inferior olive may be activated by the paracrine release of AVP from distant nuclei, such as the SCh, to control motor learning and timing. Alternatively, AVPR1a in the inferior olive may respond to other ligands (for example, OXT) found in nearby regions (72). The role of AVPR1 in the InM of male mice is less clear, but because the InM sends monosynaptic projections to the NTS (73) that is essential for blood pressure control by AVP and receiving information from the cardiovascular receptors (74), a possible coordinated control by hindbrain AVP of blood pressure and cardiovascular function.

Removed in the revised version: Although the midbrain, pons and forebrain displayed less abundant Avpr1a transcript density, they revealed further sex differences. In the midbrain and pons, the highest Avpr1a density was in the ventral tegmental decussation (vtgx) in females and dorsal raphe nucleus, interfascicular part (DRI) in males. AVP and OXT in the ventral tegmental area are known to regulate social interactions with rewarding properties.

Earlier version: Indeed, humans, as inherently social beings, show a strong inclination to affiliate and share their emotions with each other (75, 76). Sex differences in ventral tegmental AVPR1a make biological sense, as social interaction of females, specifically, with pups and, generally, with counterparts throughout their lives, have rewarding properties fundamental to maternal behavior and survival. Modulation of AVPR1a in the dorsal raphe nucleus has also been linked to social and emotional behaviors (16, 77, 78). Sexual dimorphism in AVP innervation of and AVPR1a expression in the DRI appears to imprint dimorphic social behaviors. That is, AVPR1a blockade in the lateral habenular nucleus (LHb) of males, but not females who have lesser AVP innervation of the LHb and dorsal raphe nucleus, results in reduced urine marking to unfamiliar males and ultrasonic vocalizations to unfamiliar, sexually receptive females, whereas AVPR1a blockade in the dorsal raphe nucleus of only males reduces urine marking to unfamiliar males (77).

Revised version: Although the midbrain, pons and forebrain displayed less abundant Avpr1a transcript density, they revealed further sex differences. In the midbrain and pons, the highest Avpr1a density was in the ventral tegmental decussation (vtgx) in females and dorsal raphe nucleus, interfascicular part (DRI) in males. AVP and OXT in the ventral tegmental area are known to regulate social interactions with rewarding properties. Indeed, humans, as inherently social beings, show a strong inclination to affiliate and share their emotions with each other (75, 76). Sex differences in ventral tegmental AVPR1a make biological sense, as social interaction of females, specifically, with pups and, generally, with counterparts throughout their lives, have rewarding properties fundamental to maternal behavior and survival. Modulation of AVPR1a in the dorsal raphe nucleus has also been linked to social and emotional behaviors (16, 77, 78). Sexual dimorphism in AVP innervation of and AVPR1a expression in the DRI appears to imprint dimorphic social behaviors. That is, AVPR1a blockade in the lateral habenular nucleus (LHb) of males, but not females who have lesser AVP innervation of the LHb and dorsal raphe nucleus, results in reduced urine marking to unfamiliar males and ultrasonic vocalizations to unfamiliar, sexually receptive females, whereas AVPR1a blockade in the dorsal raphe nucleus of only males reduces urine marking to unfamiliar males (77).

In the forebrain, both sexes displayed high Avpr1a transcript density in septal nuclei. The highest Avpr1a density was in the anterior commissural nucleus (AC) within the septal nuclei of females, and lateral (LSI) and medial (MS) septal nuclei of males. It is not surprising that Avpr1a transcript density was significantly greater in septal nuclei of males than females given in many rodent species males have more AVP–immunoreactive fibers in the lateral septum (46). Notably, the effects of AVP on social recognition is mediated via AVPR1a in the lateral septum (15, 26). For example, AVPR1a blockade inhibits social recognition in the rat, while AVPR1a knockout mice fail to display social recognition (15, 22, 26).

Despite Avpr1a transcript densities in other brain divisions were significantly lower, sexually dimorphic differences are worth mentioning here. In the olfactory bulb, females had the high Avpr1a density in the mitral cell layer (Mi), whereas males had high receptor expression in the olfactory tubercle (Tu). A population of AVP neurons in the olfactory bulb of the rat that play a role in social recognition via AVPR1a has been reported (79). Silencing the AVPR1a by siRNA impairs habituation/dishabituation to juvenile cues, but not to volatile odors (79). Of note, AVP is a retrograde signal that filters activation of the Mi cells in the ewe, likely through presynaptic modulation of norepinephrine or acetylcholine. The secretion of both transmitters is stimulated by AVP in the olfactory bulb (79, 80). The functional relevance of AVP signaling via AVPR1a activation in the Tu requires additional studies. There is, however, evidence in the rat that AVP via AVPR1a has, at least, an indirect impact on Tu function, as seen by a reduction in activation responding to a noxious odor of butyric acid, when the AVPR1a is blocked (81). The presence of AVPR1a in the hippocampus, thalamus, cortex, and ventricular regions are consistent with the reported effects of AVP on memory (82), emotional and reward–motivated behavior (83), blood pressure (84), blood flow and CSF production (85). Functional roles of many other nuclei shown here to express AVPR1a and not mentioned in this report are much less clear. The importance of revealing novel AVP–triggered functions by interrogating AVPR1a site–specifically, will require further investigations.

Collectively, our results provide compelling evidence of distinct and novel AVP/AVPR1a neuronal nodes in the brain. While studies on central AVP signaling and its control of blood pressure, water balance and diverse social behaviors in mammals, occupy the vast majority of the literature (3–9), we expect that this comprehensive compendium of sex–specific AVP/AVPR1a expression in the brain will deepen our understanding of the functional and neuroanatomical basis underlying old and new paradigm–shifting functions of central AVP signaling. As appears to be the case for most brain areas, the original discovery of function tends to become dogma thereby leading to an oversimplification of multiple functions of those brain areas as they interact in circuits. Finally, the approach of direct mapping of receptor expression in the brain and periphery provides the groundwork for greater discernment of new functional arrangements of ancient pituitary glycoprotein hormones and nonapeptides, such as AVP and OXT, and provide helpful pointers towards improving pharmacological interventions in disease.

Methods

Mice

Adult mice (∼3 to 4–month–old) were housed in a 12 h:12 h light : dark cycle at 22 ± 2 °C with ad libitum access to water and regular chow. All procedures were approved by the Mount Sinai Institutional Animal Care and Use Committee and are in accordance with Public Health Service and United States Department of Agriculture guidelines.

RNAscope

For RNAscope, mice were anesthetized with isoflurane (2-3% in oxygen; Baxter Healthcare, Deerfield, IL) and transcardially perfused with 0.9% heparinized saline followed by 10% Neutral Buffered Formalin (NBF). Brains were promptly extracted, post-fixed in 10% NBF for 24 h, dehydrated, and paraffin–embedded. Coronal sections were cut at 5 μm, with every tenth section mounted onto ∼60 slides with 3 sections on each slide. This method allows to cover the entire brain and eliminate the likelihood of counting the same transcript twice. Sections were air-dried overnight at room temperature and stored at 4 °C until required.

Detection of mouse Avp and Avpr (Avpr1a) was performed separately on paraffin sections using Advanced Cell Diagnostics (ACD) RNAscope 2.5 LS Reagent Kit (#322100) and two RNAscope 2.5 LS probes, namely Mm-AVP-O1 (#472268) and Mm-AVPR1a (#418068). The kidney and prefrontal cortex served as positive and negative controls for AVPR1a, respectively. As with AVP, magnocellular cells of the PVH and SON served as positive controls while the brain from AVP–knockout mouse served a negative control.

Slides were baked at 60 °C for 1 h, deparaffinized, incubated with hydrogen peroxide for 10 min at room temperature, pretreated with Target Retrieval Reagent (#322001) for 20 min at 100 °C and with Protease III for 30 min at 40 °C. Probe hybridization and signal amplification were performed as per the manufacturer’s instructions for chromogenic assays.

Following RNAscope assay, the slides were scanned at ×20 magnification and the digital image analysis was successfully validated using the CaseViewer 2.4 (3DHISTECH) software. The same software was employed to capture and prepare images for the figures in the article. Images of control tissues were taken using microscope Leica DM 1000 LED. Detection of Avp- and Avpr1a–positive cells was also performed using the QuPath-0.2.3 (University of Edinburgh, UK) software. The Atlas for the Mouse Brain in Stereotaxic Coordinates (45) was utilized to identify and manually map every nucleus, sub–nucleus or region using drawing features of the QuPath-0.2.3 software in every tenth brain section. This was followed by exhaustive counting of Avp and Avpr1a transcripts using a tag feature. Avp and Avpr1a transcript density was calculated by dividing the absolute numbers by the total area (µm2, ImageJ) of every nucleus, sub–nucleus or region. Photomicrographs were prepared using Photoshop CS5 (Adobe Systems) only to adjust brightness, contrast and sharpness, to remove artifacts (i.e., obscuring bubbles), and to make composite plates.

Quantitation, Validation and Statistical Analysis

Data were analyzed by Student’s t-test and two-way analysis of variance (ANOVA) followed by Tukey’s multiple comparisons tests using GraphPad Prism 10.2.2 version (La Jolla, CA). Significance was set at P < 0.05. P values are shown.